```python
# Initialize Otter
import otter
grader = otter.Notebook("dsc495_0019_r3.ipynb")
```

# Week 3 Review

## Due: Friday, February 4, 2022 @ 11:59pm

In this assignment we will review the concepts and topics that we covered in Week 3.

I would like you to attempt each level of the assignment. They get progressively harder. To get a **comprehensive review** of last week's content, you must complete all levels.

**Note:** Try not to delete the instructions of the assignment.

In the markdown cell below enter your name, section, and the date.

**Name:** 

**Section:** 

**Date:**

## Pandas Overview

[Pandas](https://pandas.pydata.org/) is one of the most widely used Python libraries in data science. In this lab, you will learn commonly used data wrangling operations/tools in Pandas. We aim to give you familiarity with:

* Slicing data frames (i.e. selecting rows and columns)
* Filtering data (using boolean arrays)

Run the cell below.

```python
import pandas as pd
import numpy as np
```

## Basic Manipulations

### Demographics

Let's take a look at some demographic information about the counties in North Carolina.

```python
demo = pd.read_csv('data/county_demographics.csv')
demo.head()
```

The `.info()` and `.describe()` methods for dataframes give you information about the data. 

Run each command in the cells below.

```python
# Use this cell for .info
...
```

```python
# Use this cell for .describe
...
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Level I

**Question 1.** Let's do some comparisons of North Carolina to the rest of the country. Filter the data and save a new dataframe with only the NC data. Name this dataframe `demo_nc`.

```python
demo_nc = ...
demo_nc.head()
```

```python
grader.check("q1")
```

## Level II

**Question 2.** What is the median percent of county residents with a Bachelor's Degree of higher? Save this value to `bach_degree_percent`.

**Note:** Be careful about the puncutation in the column names.

```python
bach_degree_percent = ...
```

```python
grader.check("q2")
```

**Question 3.**  Which counties have at least the median (i.e. the value you caluclated in **Question 2.**)? Save this to a dataframe named `demo_nc_bach`.

```python
demo_nc_bach = ...
demo_nc_bach.head()
```

```python
grader.check("q3")
```

## Level III

<!-- BEGIN QUESTION -->

**Question 4.** Suppose we are interested in the percentage of county residents that complete high school and the percentage of county residents that complete a Bachelor's Degree. Make a dataframe that countains the name of the county, the state, the percentage of county residents that complete high school and the percentage of county residents that complete a Bachelor's Degree.

**Note:** You do not need to save this dataframe to an object. Just display it.

**Hint:** You can subset mulitple columns by passing a list object to the square braces, `[[list]]`.

```python
demo_nc[["County", "State", "Education.Bachelor's Degree or Higher", "Education.High School or Higher"]]
```

<!-- END QUESTION -->

## Level IV

**Question 5.** Save the dataframe from **Question 4.** to `demo_nc_ed`. Change the name of the education columns to `Bachelor`, `High School`.

```python
demo_nc_ed = ...
demo_nc_ed
```

```python
grader.check("q5")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> demo_nc.shape
(100, 52)
```


Tests for q2 (points: None)
```
>>> bach_degree_percent
17.9
```


Tests for q3 (points: None)
```
>>> isinstance(demo_nc_bach, pd.core.frame.DataFrame)
True
```
```
>>> demo_nc_bach.shape
(51, 52)
```


Tests for q5 (points: None)
```
>>> demo_nc_ed.shape
(100, 4)
```
```
>>> demo_nc_ed.columns
Index(['County', 'State', 'Bachelor', 'High School'], dtype='object')
```
