```python
# Initialize Otter
import otter
grader = otter.Notebook("dsc495_0019_a3.ipynb")
```

# Assignment 3

## Due: Tuesday, February 1, 2022 @ 11:59pm

In this assignment we will get more familiar with basic Python, including the `NumPy` operations, manipulating `DataFrame`s, `for` loops, and `if` `elif` `else` statements.

I would like you to attempt each level of the assignment. They get progressively harder. To get **full credit** on this assignment, you must complete the first three levels (Levels I, II, and III). Levels IV and V give you a chance to work with some more advanced printing freatures in Python.

**Note:** Try not to delete the instructions of the assignment.

In the markdown cell below enter your name, section, and the date.

**Name:** 

**Section:** 

**Date:**

## Level I

### `NumPy`

Very often in data science, we want to work with many numbers that are evenly spaced within some range.

**Question 1.** Import `numpy` as `np` and then use `np.arange` to create an array with the multiples of 99 from 0 up to (**and including**) 9999. So its elements are 0, 99, 198, 297, etc.

```python
# Imoprt the pandas package
import ... as np

# Create the numpy array
multiples_of_99 = ...
multiples_of_99
```

```python
grader.check("q1")
```

### Working with Array Elements

Let's work with a more interesting dataset.  The next cell creates an array called `population_amounts` that includes estimated world populations in every year from **1960** to roughly the present. The estimates come from the US Census Bureau website.

Rather than type in the data manually, we've loaded a file on your computer called `world_population.csv`.

**Question 2.** Import `pandas` as `pd` and then load the `world_population.csv` from the data folder. Save it to `population_amounts`.

```python
# Imoprt the pandas package
import ... as pd

# Load the world_population.csv file
population_amounts = ...('data/world_population.csv')
population_amounts
```

```python
grader.check("q2")
```

We can convert a `pandas` `Series` (i.e. a column from a `pandas` `DataFrame`) into a `NumPy` array using the `.to_numpy()` method. Click [here]() to read about the `.to_numpy()` method.

**Question 3.** Convert the `Population` column from the `population_amounts` `DataFrame`, into a `NumPy` array. Save it to `population_array`.

```python
population_array = ...
population_array
```

```python
grader.check("q3")
```

Run the cell below.

```python
population_array[0]
```

The value of that expression is the number 3032156070 (around 3 billion), because that's the first element in the array `population_array`.

```python
# The 13th element in the array is the population in 1972 (which is 1960 + 12).
population_1972 = population_array[12]
population_1972
```

**Question 4.** Set `population_1983` to the world population in 1983, by getting the appropriate element from `population_array`.

```python
population_1983 = ...
population_1983
```

```python
grader.check("q4")
```

### Doing Something to Every Element of an Array

Arrays are primarily useful for doing the same operation many times. Arithmetic works elementwise on arrays, meaning that if you perform an arithmetic operation (like subtraction, division, etc.) on an array, Python will do the operation to every element of the array individually and return an array of all of the results. For example, you can divide all the population numbers by 1 billion to get numbers in billions.

**Question 5.** Divide all the population numbers by 1 billion to get numbers in billions. Save the results to `population_in_billions`. 

```python
population_in_billions = ...
population_in_billions
```

```python
grader.check("q5")
```

## Level II

#### Logarithms

Orders of magnitude quantify how big a number is by representing it as the power of another number (for example, representing 100 as $10^2$). One way to do this is by using the logarithm function. The logarithm (base 10) of a number increases by 1 every time we multiply the number by 10. It's like a measure of how many decimal digits the number has, or how big it is in orders of magnitude.

Suppose we were interested in how big was the population in **orders of magnitude** in each year? 

We could try to answer our question like this, using the `log10` function from the `math` module.

```python
import math 

population_1960_magnitude = math.log10(population_array[0])
population_1961_magnitude = math.log10(population_array[1])
population_1962_magnitude = math.log10(population_array[2])
population_1963_magnitude = math.log10(population_array[3])

population_1960_magnitude, population_1961_magnitude, population_1962_magnitude, population_1963_magnitude
```

But this is tedious and doesn't really take advantage of the fact that we are using Python.

Instead, `NumPy` provides its own version of `log10` that takes the logarithm of each element of an array.  It takes a single array of numbers as its argument.  It returns an array of the same length, where the first element of the result is the logarithm of the first element of the argument, and so on.

The documentation can be found by clicking [here](https://numpy.org/doc/stable/reference/generated/numpy.log10.html).

**Question 6.** Use `np.log10` to compute the logarithms of the world population in every year.  Give the result (an array of 66 numbers) the name `population_magnitudes`. 

```python
population_magnitudes = ...
population_magnitudes
```

```python
grader.check("q6")
```

What you just did is called **elementwise** application of `np.log10`, since `np.log10` operates separately on each element of the array that it's called on.

## Level III

#### Powers

The powers of 2 $\left(2^0 = 1, 2^1 = 2, 2^2 = 4, \text{ etc} \right)$ arise frequently in computer science. For example, you may have noticed that storage on smartphones come in powers of 2, like 16 GB, 32 GB, or 64 GB.

**Question 7.** Use `np.arange` and the exponentiation operator `**` to compute the first 30 powers of 2, starting from $2^0$. Save this to `powers_of_2`.

**Hint:** Part of your solution will involve `np.arange`, but your array shouldn't have more than 30 elements.

**Warning:** `np.arange(1, 2**30, 1)` creates an array with $2^{30}$ elements and **will crash your kernel**.

```python
powers_of_2 = ...
powers_of_2
```

```python
grader.check("q7")
```

## `pandas` Operations

### CEO Incomes

Let's take a look at the 2019 compensation of North Carolina CEOs. The following table contains information on the top 61 highest paid CEOs for companies headquartered in North Carolina. The data was compiled from the [AFL-CIO Executive Paywatch database](https://aflcio.org/paywatch/highest-paid-ceos), and ultimately came from [filings](https://www.sec.gov/answers/proxyhtf.htm) mandated by the SEC from all publicly-traded companies.

We've copied the raw data from the AFL-CIO database into a file called `nc-ceo-pay.csv`.

```python
raw_compensation = pd.read_csv('data/nc-ceo-pay.csv')
raw_compensation
```

**Question 9.** Select all the rows for companies in Raleigh. Save this to a `DataFrame` named `raleigh_ceo_salaries`.

```python
raleigh_ceo_salaries = ...
raleigh_ceo_salaries
```

```python
grader.check("q9")
```

**Question 10.** Which CEO has the highest **salary** in the Raleigh-Durham area? Assign the name of the CEO (as it is spelled in the dataset) to `top_salary`.

```python
top_salary = ...
```

```python
grader.check("q10")
```

**Question 11.** Which CEO has the highest **bonus** in the Raleigh-Durham area? Assign the name of the CEO (as it is spelled in the dataset) to `top_bonus`.

```python
top_bonus = ...
```

```python
grader.check("q11")
```

### Counting Words

#### The Gettysburg Address

The Gettysburg Address is a speech that U.S. President Abraham Lincoln delivered during the American Civil War at the dedication of the Soldiers' National Cemetery in Gettysburg, Pennsylvania, on the afternoon of November 19, 1863, four and a half months after the Union armies defeated those of the Confederacy at the Battle of Gettysburg. It is one of the best-known speeches in American history.

**Source:** [Wikipedia](https://en.wikipedia.org/wiki/Gettysburg_Address)

Run the cell below to load the text file and display the first sentence.

**Source:** Prepared by Gerald Murphy (The Cleveland Free-Net - aa300) Distributed by the Cybercasting Services Division of the National Public Telecomputing Network (NPTN).

```python
with open('data/gettysburg-address.txt', 'r') as f:
    lines = f.read()
print(lines[:175])
```

The file was read in and saved as a single string.

Run the cell below to see the full string. You should notice the `\n` character sequence, the standard escape sequence for end of line.

```python
lines[:]
```

**Question 12.** Change all the letters to lowercase using a string method. Then, replace the `\n` character sequence with a `' '` (this is a space) using a string method. Then we can split the string at the spaces using `.split` - this will reurn a list. Save this to the object `lines_with_spaces`.

**Hint:** You can do this in one line. That being said, feel free to use multiple lines if it will be easier for you. 

```python
lines_with_spaces = ...

# Print the first 10 words
lines_with_spaces[0:6]
```

```python
grader.check("q12")
```

**Question 13.** Use a loop to count the number of words in the speech (use the object `lines_with_spaces`). Save the number of words in `word_count`. 

**Note:** To receive full credit on this question you must use a loop.

```python
# Set counter to 0
word_count = ...

# Loop over the list of words and update counter
for ...
word_count
```

```python
grader.check("q13")
```

**Question 14.** Use a loop and a `if` statement to count the number of vowels (a, e, i, o, and u) in the speech (use the object `lines`). Save the number of words in `vowel_count`. Make sure you don't count the `\n` character sequence.

**Note:** To receive full credit on this question you must use a loop.

```python
# Set counter to 0
word_count = ...

# Loop over the list of words and update counter
for ...
vowel_count
```

```python
grader.check("q14")
```

**Question 15.** Use a loop and a `if` statement to count the number of consonants, and spaces in the speech (use the object `lines`). Make sure you don't count the `\n` character sequence. Then print the following message:

    The number of vowels is: ####
    The number of consonants is: ####
    The number of spaces are: ####
    
**Note:** Use the `vowel_count` variable from **Question 14.** to store the number of vowels, use `consonant_count` to store the number of consonants, and use `spaces_count` to store the number of spaces.

```python
...

vowel_count = ...
consonant_count = ...
spaces_count = ...

print("The number of vowels: ", ...)
print("The number of consonants: ", ...) 
print("The number of spaces: ", ...)
```

## Level IV

### Numerical Summaries of CEO Incomes

Let's take another look at the 2019 compensation of North Carolina CEOs. This time we are interested in the mean and median total pay.

Run the cell below to load the dataset into a different `DataFrame` named `nc_ceo_totalpay`.

```python
nc_ceo_totalpay = pd.read_csv('data/nc-ceo-pay.csv')
nc_ceo_totalpay
```

The data type first item in the `Total Pay` column is a `str`. This means we won't be able to perform any arithmetic on it.

Runt the cell below to see.

```python
type(nc_ceo_totalpay['Total Pay'][0])
```

 Another clue that could have lead you to this discovery is the fact that there are characters in the observation.
 
 Run the cell below to see.

```python
nc_ceo_totalpay['Total Pay'][0]
```

We need to remove all the characters that are not numbers (i.e. the \$ and ,), then convert all the elements from strings (`str`) to numbers (`int`). After that we can find the mean and median.

There are several string methods that can be used on `pandas` `Series`. They can be found [here](https://pandas.pydata.org/docs/reference/api/pandas.Series.str.html). There are also `Series` methods that can deal with data types. Those can be found [here](https://pandas.pydata.org/docs/reference/series.html).

**Question 16.** Find the median, mean, and standard devitaion of the total pay for using the `nc_ceo_totalpay` `DataFrame`. Store the values in `nc_ceo_mean` and  `nc_ceo_median`. Round each value to remove all the numbers after the decimal.
    
**Note:** You **do not** need to use a loop. 

**Remember:** You can use functions from the `NumPy` library to do *elementwise* arithmetic.

```python
nc_ceo_mean = ...
nc_ceo_median = ...
nc_ceo_sd = ...
nc_ceo_mean, nc_ceo_median
```

```python
grader.check("q16")
```

## Level V

Now we want to visualize the distribution of total pay. First we need to import a plotting module.

Run the cell below to load `matplotlib`.

```python
import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")
%matplotlib inline
```

<!-- BEGIN QUESTION -->

**Question 17.** Make a histogram of the values in the `Total Pay` column. 

```python
# Hide gridlines
plt.grid(False)

# Add title
plt.title("Distributions of NC CEO Total Pay")

# Add x label
plt.xlabel("Total Pay (in millions of dollars)")

# Add y label
plt.ylabel("Frequency")

# Plot histogram 
...;
```

<!-- END QUESTION -->

Based on our histogram it looks like there are some CEOs that made more than 10 million dollars in total pay.

**Question 18.** Create a `DataFrame` that contains the name, company, and total pay for the CEOs that made at least 10 million dollars in 2019. Save this to `nc_ceo_10million`, with column names `name`, `company`, `total_pay`.

**Hint:** You can create a data frame by specifying the columns and values using a [dictionary](https://www.geeksforgeeks.org/python-dictionary/) as shown [here](https://www.geeksforgeeks.org/creating-a-dataframe-from-pandas-series/) in the **Code \#2: Creating Dataframe from Series** example. The keys of the [dictionary](https://www.geeksforgeeks.org/python-dictionary/) are the column names, and the values of the [dictionary](https://www.geeksforgeeks.org/python-dictionary/) are lists containing the row entries.

**Note:** You can use multiple lines of code to complete this question.

```python
...
nc_ceo_10million
```

```python
grader.check("q18")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> isinstance(multiples_of_99, np.ndarray)
True
```
```
>>> len(multiples_of_99)
102
```


Tests for q10 (points: None)
```
>>> isinstance(top_salary, str)
True
```


Tests for q11 (points: None)
```
>>> isinstance(top_bonus, str)
True
```


Tests for q12 (points: None)
```
>>> "\n" not in lines_with_spaces
True
```
```
>>> string = ' '.join([str(item) for item in lines_with_spaces])
>>> string.islower()
True
```


Tests for q13 (points: 0.5)
```
>>> isinstance(word_count, int)
True
```


Tests for q14 (points: 0.5)
```
>>> isinstance(vowel_count, int)
True
```


Tests for q16 (points: 0)
```
>>> nc_ceo_mean
6332532.0
```
```
>>> nc_ceo_median
4850813.0
```


Tests for q18 (points: 0)
```
>>> nc_ceo_10million.shape
(12, 3)
```
```
>>> nc_ceo_10million.columns
Index(['name', 'company', 'total_pay'], dtype='object')
```
```
>>> np.sum(int(nc_ceo_10million['total_pay']))
```


Tests for q2 (points: None)
```
>>> isinstance(population_amounts, pd.core.frame.DataFrame)
True
```
```
>>> population_amounts.shape
(61, 1)
```
```
>>> np.sum(population_amounts['Population'])
323331061383
```


Tests for q3 (points: None)
```
>>> isinstance(population_array, np.ndarray)
True
```
```
>>> len(population_array)
61
```
```
>>> np.sum(population_array)
323331061383
```


Tests for q4 (points: None)
```
>>> population_1983
4674330282
```


Tests for q5 (points: None)
```
>>> isinstance(population_in_billions, np.ndarray)
True
```
```
>>> len(population_in_billions)
61
```
```
>>> sum(population_in_billions)
323.331061383
```


Tests for q6 (points: None)
```
>>> isinstance(population_magnitudes, np.ndarray)
True
```
```
>>> len(population_magnitudes)
61
```
```
>>> sum(population_magnitudes)
592.1786136553848
```


Tests for q7 (points: None)
```
>>> isinstance(powers_of_2, np.ndarray)
True
```
```
>>> sum(powers_of_2)
1073741823
```


Tests for q9 (points: None)
```
>>> raleigh_ceo_salaries.shape
(8, 12)
```
```
>>> isinstance(raleigh_ceo_salaries, pd.core.frame.DataFrame)
True
```
