```python
# Initialize Otter
import otter
grader = otter.Notebook("dsc495_0019_r1.ipynb")
```

# Week 1 Review

## Due: Thursday, January 20, 2022 @ 11:59pm

In this assignment we will review the concepts and topics that we covered in Week 1.

I would like you to attempt each level of the assignment. They get progressively harder. To get **full credit** on this assignment, you must complete the first two levels (Levels I and II). Level III gives you a chance to work with some more advanced printing freatures in Python.

**Note:** Try not to delete the instructions of the assignment.

In the markdown cell below enter your name, section, and the date.

**Name:** 

**Section:** 

**Date:**

<!-- BEGIN QUESTION -->

## Level I

**Question 1.** Write code in the cell below that prints out the following statements.

```
Course Name: R/Python for Data Science
Course Code: DSC 495
Sections: 001 009
```

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 2.** In our first lesson we listed the seven steps of the Data Science Methodolgy. 

1. Problem Understanding Phase
2. Data Preparation Phase
3. Exploratory Data Analysis Phase
4. Setup Phase
5. Modeling Phase
6. Evaluation Phase
7. Deployment Phase

Print the steps that require us to use Python/R.

```python
...
```

<!-- END QUESTION -->

## Level II

**Question 3.** Raw data from data repositories is usually ready for use in the form that it is downloaded. Instead, it needs to be cleaned for analysis. From the list below choose the step of the Data Science Methodolgy that describes the process mentioned in the previous sentence.

* Problem Understanding Phase
* Data Preparation Phase
* Exploratory Data Analysis Phase
* Setup Phase
* Modeling Phase
* Evaluation Phase
* Deployment Phase

Assign to the correct answer choice to the variable `q3_answer`. 

**Note:** For example, `q3_answer = "Your Answer"`.

```python
q3_answer = "...""
q3_answer
```

```python
grader.check("q3")
```

## Level III

### Lists 

In Python lists are one of the built-in data types used to store multiple items in a single variable. List items are ordered by index and the first index is `[0]`, the second item has index `[1]`, etc. 

Lists are defined using `[ ]` and can be assigned to a variable name. For example

`my_list = ["Data", "Science", "Academy"]`

has three items `[0]`, `[1]`, and `[2]`. Each element in the list is a string. Strings in python are surrounded by either single quotation marks, or double quotation marks.

**Question 4.** Create a list named `sections` that contains the section codes (as strings) for this course. Then use the `print()` function to print the list.

**Note:** The course section numbers are 001 and 009.

```python
sections = ...
print(...)
```

```python
grader.check("q4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for q3 (points: None)
```
>>> isinstance(q3_answer, str)
True
```
```
>>> q3_answer == "Data Preparation Phase"
True
```


Tests for q4 (points: 0)
```
>>> isinstance(sections, list)
True
```
```
>>> len(sections)
2
```
```
>>> sections[0]
'001'
```
```
>>> sections[1]
'009'
```
