```python
# Initialize Otter
import otter
grader = otter.Notebook("dsc495_0019_a2.ipynb")
```

# Assignment 2

## Due: Tuesday, January 25, 2022 @ 11:59pm

In this assignment we will get more familiar with basic Python, including the expressions, functions and variable assignment.

I would like you to attempt each level of the assignment. They get progressively harder. To get **full credit** on this assignment, you must complete the first three levels (Levels I, II, and III). Levels IV and V gives you a chance to work with some more advanced printing freatures in Python.

**Note:** Try not to delete the instructions of the assignment.

In the markdown cell below enter your name, section, and the date.

**Name:** 

**Section:** 

**Date:**

## Level I

### The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code. 

Next to every code cell, you'll see some text that says `In [...]`. Before you run the cell, you'll see `In [ ]`. When the cell is running, you'll see `In [*]`. If you see an asterisk (\*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `In [1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection. If this happens, try the following steps:

1. At the top of your screen, click **Kernel**, then **Interrupt Kernel**. Trying running your code again.

2. If that doesn't help, click **Kernel**, then **Restart Kernel**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.

3. If that doesn't help, restart your server. This is very rarely needed, but it's good to know how to do just in case. First, save your work by clicking **File** at the top left of your screen, then **Save Notebook**. Next, click **File** again and choose **Hub Control Panel** and select **Stop My Server** to shut it down. Then, after a few moments select **Start My Server** to start it back up. Then, navigate back to the notebook you were working on. You'll still have to run your code cells again.

### Numbers

Quantitative information arises everywhere in data science. In addition to representing commands to print out lines, expressions can represent numbers and methods of combining numbers. The expression `3.14159` evaluates to the number 3.14159. 

Run the cell and see.

```python
3.14159
```

Notice that we didn't have to `print`. When you run a notebook cell, if the last line has a value, then Jupyter helpfully prints out that value for you. However, it won't print out prior lines automatically.

```python
print(3.14159)
4
3
2
1
```

Above, you should see that 1 is the value of the last expression, 3.14159 is printed, but 4, 3, and 2 are all lost forever because they were neither printed nor last.

**Note:** You don't want to print everything all the time anyway.

### Arithmetic

The line in the next cell subtracts.  Its value is what you'd expect.  

Run it.

```python
3.75-1.5
```

Many basic arithmetic operations are built into Python. 

| Expression Type | Operator | Example | Value |
| --------------- | -------- | ------- | ----- |
| Addition        | $+$      | $2+3$   | $5$   |
| Subtraction     | $-$      | $2-3$   | $-1$  |
| Multiplication  | $*$      | $2*3$   | $6$   |
| Division        | $/$      | $3/2$   | $1.5$ |
| Remainder       | $%$      | $3%2$   | $1$   |
| Exponentiation  | $**$     | $3**2$  | $9$   |

The common operator that differs from typical math notation is `**`, which raises one number to the power of the other. So, `3**2` stands for $3^2$ and evaluates to 9. 

The order of operations is the same as what you learned in elementary school, and Python also has parentheses. For example, compare the outputs of the cells below. The second cell uses parentheses for a happy new year.

```python
6+6*5-6*3**2*2**3/4*7
```

```python
6+(6*5-(6*3))**2*((2**3)/4*7)
```

**Question 1.** Write a Python expression in this next cell that's equal to 

$$5 \times \left(3 \frac{10}{11}\right)-50\frac{1}{3}+2^{.5 \times 22}-\frac{7}{33}+5$$  

That's five times three and ten elevenths, minus fifty and a third, plus two to the power of half twenty-two, minus seven thirty-thirds plus five.  By "$3 \frac{10}{11}$" we mean $3+\frac{10}{11}$, not $3 \times \frac{10}{11}$.

Replace the ellipses (`...`) with your expression.  Try to use parentheses only when necessary.

**Hint:** The correct output should be a familiar number.

```python
...
```

```python
grader.check("q1")
```

### Names

In natural language, we have terminology that lets us quickly reference very complicated concepts.  We don't say, "That's a large mammal with brown fur and sharp teeth!"  Instead, we just say, "Bear!"

In Python, we do this with **assignment statements**. An assignment statement has a name on the left side of an `=` sign and an expression to be evaluated on the right.

```python
three_thats_the_magic_number = 1+1+1
```

When you ran the previous cell, Python first computed the value of the expression on the right-hand side, `1+1+1`, which is the number 3.  Then it assigned that value to the name `three_thats_the_magic_number`.  At that point, the code in the cell was done running.

After you run that cell, the value 3 is bound to the name `three_thats_the_magic_number`:

```python
three_thats_the_magic_number
```

A common pattern in Jupyter notebooks is to assign a value to a name and then immediately evaluate the name in the last line in the cell so that the value is displayed as output. 

```python
close_to_pi = 355/113
close_to_pi
```

Another common pattern is that a series of lines in a single cell will build up a complex computation in stages, naming the intermediate results.

```python
weekly_salary = 925.75
monthly_salary = 4*weekly_salary
number_of_months_in_a_year = 12
yearly_salary = number_of_months_in_a_year * monthly_salary
yearly_salary
```

Names in Python can have letters (upper- and lower-case letters are both okay and count as different letters), underscores, and numbers.  The first character can't be a number (otherwise a name might look like a number).  And names can't contain spaces, since spaces are used to separate pieces of code from each other.

Other than those rules, what you name something doesn't matter to Python. For example, this cell does the same thing as the above cell, except everything has a different name:

```python
a = 925.75
b = 4*a
c = 12
d = c*b
d
```

**However**, names are very important for making your code readable to yourself and others.  The cell above is shorter, but it's totally useless without an explanation of what it does.

**Question 2.** Assign the name `seconds_in_a_decade` to the number of seconds between midnight January 1, 2010 and midnight January 1, 2020. 

**Note:** There are two leap years in this span of a decade. A non-leap year has 365 days and a leap year has 366 days.

```python
seconds_in_a_decade = ...

# The last line in this cell will print the value
# you've given to seconds_in_a_decade when you run it.  
# You don't need to change this.
seconds_in_a_decade
```

```python
grader.check("q2")
```

### Comments

You may have noticed these lines in the cell in which you answered **Question 2.**

    # The last line in this cell will print the value
    # you've given to seconds_in_a_decade when you run it.  
    # You don't need to change this.
    
This is called a **comment**. It doesn't make anything happen in Python; Python ignores anything on a line after a `#`.  Instead, it's there to communicate something about the code to you, the human reader. Comments are extremely useful. 

## Level II

### Expressions

The two building blocks of Python code are *expressions* and *statements*.  An **expression** is a piece of code that

* is self-contained, meaning it would make sense to write it on a line by itself, and
* usually evaluates to a value.


Here are two expressions that both evaluate to 3

* `3`
* `5-2`
    
One important type of expression is the **call expression**. A call expression begins with the name of a function and is followed by the argument(s) of that function in parentheses. The function returns some value, based on its arguments. Some important mathematical functions are listed below.

| Function | Description                                                   |
|----------|---------------------------------------------------------------|
| `abs`    | Returns the absolute value of its argument                    |
| `max`    | Returns the maximum of all its arguments                      |
| `min`    | Returns the minimum of all its arguments                      |
| `pow`    | Raises its first argument to the power of its second argument |
| `round`  | Rounds its argument to the nearest integer                    |

Here are two call expressions that both evaluate to 3.

* `abs(2-5)`

* `max(round(2.8), min(pow(2, 10), -1*pow(2, 10)))`

The expression `5-2` and the two call expressions given above are examples of **compound expressions**, meaning that they are actually combinations of several smaller expressions. `5-2` combines the expressions `5` and `2` by subtraction.  In this case, `5` and `2` are called **subexpressions** because they're expressions that are part of a larger expression.

A **statement** is a whole line of code.  Some statements are just expressions.  The expressions listed above are examples.

Other statements *make something happen* rather than *having a value*. For example, an **assignment statement** assigns a value to a name. 

A good way to think about this is that we're **evaluating the right-hand side** of the equals sign and **assigning it to the left-hand side**. Here are some assignment statements:
    
* `height = 65`

* `absolute_height_difference = abs(height-72)`

An important idea in programming is that large, interesting things can be built by combining many simple, uninteresting things. The key to understanding a complicated piece of code is breaking it down into its simple components.

For example, a lot is going on in the last statement above, but it's really just a combination of a few things.


* `height` is a name expression assignned in the first line

* `72` is a number expression whose value is `72`

* `height-72` is an arithmetic expression whose value is `-7`

* `abs()` is a name expression that is a function whose value is the absolute value of the arithmetic epxression `height-72`

* `absolute_height_difference` is part of an assignment statement that assigns the value of `7` to the variable name

Run the cell below.

```python
height = 65
absolute_height_difference = abs(height-72)
absolute_height_difference
```

### Nested Expressions

Function calls and arithmetic expressions can themselves contain expressions.  You saw an example in the previosu cell `abs(height-72)` has a name expression and a number expression, in a subtraction expression, in a function call expression.

Suppose we are interested in heights that are very unusual. We'll say that a height is unusual to the extent that it's far away from the average human height. [An estimate](http://www.wecare4eyes.com/averageemployeeheights.htm) of the average adult human height (averaging, we hope, over all humans on Earth today) is 1.688 meters. Muggsy Bogues played for the Charlotte Hornets. He is shortest player in NBA history at 1.6002 meters tall, therefore his height is $|1.6002-1.688|$, or $0.0878$ meters away from the average.

Here's how we'd write that in one line of Python code:

```python
round(abs(1.6002-1.688), 4)
```

What's going on here?  `round` takes two arguments, and `abs` takes one argument. The stuff inside the parentheses after `abs` is all part of that **single argument**.  Specifically, the argument is the value of the expression `1.6002-1.688`.  The value of that expression is `-0.08779999999999988`.  That value is the argument to `abs`.  The absolute value of that is `0.08779999999999988`. So `0.08779999999999988` is the value of the full expression `abs(1.6002-1.688)`. 

The value `0.08779999999999988` is the first argument of the `round` function and `4` is the second argument. This value `0.08779999999999988`will be rounded to four decimal places to `0.0878`.

Picture simplifying the expression in several steps:

1. `round(abs(1.6002-1.688), 4)`

2. `round(abs(-0.08779999999999988), 4)`

3. `round(0.08779999999999988, 4)`

4. `0.0878`

In fact, that's basically what Python does to compute the value of the expression.

**Question 3.** Given the heights of three players from the Charlotte Hornets, write an expression that computes the smallest difference between any of the three heights rounded to four decimal places. Your expression shouldn't have any numbers in it, only function calls and the names `lamelo`, `hayward`, and `bridges`. Give the value of your expression the name `min_height_difference`.

```python
# The three players' heights, in meters:
lamelo_m = 1.98   # LaMelo Ball is 6'6"
hayward_m = 2.01  # Gordon Hayward is 6'7"
bridges_m = 2.10  # Miles Bridges is 6'9"
             
# We'd like to look at all 3 pairs of heights, 
# compute the absolute difference between each pair, 
# and then find the smallest of those 3 absolute differences.  

# This is left to you. 
# If you're stuck, try computing the value for each step of the process 
# (like the difference between LaMelo's heigh and Hayward's height) 
# on a separate line and giving it a name (like lamelo_hayward_height_diff)
min_height_difference = ...

# Again, we've written this here so that
# the distance you compute will get printed 
# when you run this cell.
min_height_difference
```

```python
grader.check("q3")
```

## Level III

### Importing code

Most programming involves work that is very similar to work that has been done before. Since writing code is time-consuming, it's good to rely on others' published code when you can. Rather than copy-pasting, Python allows us to **import modules**. A module is a file with Python code that has defined variables and functions. By importing a module, we are able to use its code in our own notebook.

### `math`

Python includes many useful modules that are just an import away.  We'll look at the `math` module as a first example. The `math` module is extremely useful in computing mathematical expressions in Python. 

Suppose we want to very accurately compute the area of a circle with a radius of 5 meters. For that, we need the constant $\pi$, which is roughly 3.14159.  Conveniently, the `math` module has `pi` defined for us.

Run the cell below.

```python
import math

radius = 5
area_of_circle = radius**2 * math.pi
area_of_circle
```

In the code above, the line `import math` imports the math module. This statement creates a module and then assigns the name `math` to that module. We are now able to access any variables or functions defined within `math` by typing the name of the module followed by a dot, then followed by the name of the variable or function we want.

**Question 4.** The module `math` also provides the name `e` for the base of the natural logarithm, which is roughly 2.71.  Compute $\displaystyle e^{\pi}-\pi$, giving it the name `near_twenty`.

```python
near_twenty = ...
near_twenty
```

```python
grader.check("q4")
```

For your reference, the cells below demonstrate some more examples of functions from the `math` module.

Notice how different functions take in different numbers of arguments. Often, the [documentation](https://docs.python.org/3/library/math.html) of the module will provide information on how many arguments are required for each function.

Run the cell below.

```python
# Calculating logarithms (the logarithm of 8 in base 2)
# The result is 3 because 2 to the power of 3 is 8
math.log(8, 2)
```

Run the cell below.

```python
# Calculating square roots
math.sqrt(5)
```

There are various ways to import and access code from outside sources. The method we used above — `import <module_name>` — imports the entire module and requires that we use `<module_name>.<name>` to access its code. 

We can also import a specific constant or function instead of the entire module. Notice that you don't have to use the module name beforehand to reference that particular value. However, you do have to be careful about reassigning the names of the constants or functions to other values.

Run the cell below.

```python
# Importing just cos and pi from math
# We don't have to use `math.` in front of cos or pi

from math import cos, pi
print(cos(pi))

# We do have to use it in front of other functions from math, though
math.log(pi)
```

Or we can import every function and value from the entire module.

Run the cell below.

```python
# Lastly, we can import everything from math using the *
# Once again, we don't have to use 'math.' beforehand 

from math import *
log(pi)
```

`math` provides the name `sin` for the sine function.  Having imported `math` already, we can write `sin(3)` to compute the sine of 3.

**Note:** This sine function considers its argument to be in [radians](https://en.wikipedia.org/wiki/Radian), not degrees. 180 degrees are equivalent to $\pi$ radians.

**Question 5.** A $\frac{\pi}{4}$ radian (45-degree) angle forms a right triangle with equal base and height, pictured below.  If the hypotenuse (the radius of the circle in the picture) is 1, then the height is $\sin\left(\frac{\pi}{4}\right)$.  Compute that value using `sin` and `pi` from the `math` module.  Give the result the name `sine_of_pi_over_four`.

<center><img src="http://mathworld.wolfram.com/images/eps-gif/TrigonometryAnglesPi4_1000.gif"></center>

**Source:** [Wolfram MathWorld](http://mathworld.wolfram.com/images/eps-gif/TrigonometryAnglesPi4_1000.gif)

```python
sine_of_pi_over_four = ...
sine_of_pi_over_four
```

```python
grader.check("q5")
```

### `pandas`

#### Indeed

Indeed calculates the percentage change in seasonally-adjusted job postings since February 1, 2020, using a 7-day trailing average. February 1, 2020, is the pre-pandemic baseline. Indeed seasonally adjusts each series based on historical patterns in 2017, 2018, and 2019. Each series, including the national trend, occupational sectors, and sub-national geographies, is seasonally adjusted separately. Indeed switched to this new methodology in January 2021 and now reports all historical data using this new methodology. Historical numbers have been revised and may differ significantly from originally reported values. The new methodology applies a detrended seasonal adjustment factor to the percentage change in job postings. For more information, see Frequently Asked Questions regarding Indeed Data.

The data file `indeed_data.csv` contains information for Raleigh-Cary (`ral_cary`), Charlotte-Concord-Gastonia (`char_con_gast`), and Greensboro-High Point (`gso_hp`).

**Question 6.** Import `pandas` as `pd` and then use `pd.read_csv` to read load a `.csv` file into a `pandas` `DataFrame` named `indeed`. 

**Note:** The file `indeed_data.csv` is located in the `data` folder.

```python
# Imoprt the pandas package
import ... as pd

# Load the indeed_data.csv file
indeed = ...
indeed.head()
```

```python
grader.check("q6")
```

**Question 7.** Use the `.set_index()` method to change the index to the `observation_date`. Store this DataFrame in `indeed_date_idx`. Click [here](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.set_index.html) to read the documentation.

```python
indeed_date_idx = ...
indeed_date_idx.head()
```

```python
grader.check("q7")
```

**Question 8.** Select all the rows from the `indeed_date_idx` `DataFrame` that correspond to February 1, 2020 through February 29, 2020. Store these rows in a `DataFrame` named `indeed_feb2020`.

```python
indeed_feb2020 = ...
indeed_feb2020.head()
```

```python
grader.check("q8")
```

**Question 9.** Select all the rows from the `indeed_feb2020` `DataFrame` where the value in the `char_con_gast` column is positive. Store these rows in a `DataFrame` named `indeed_feb2020_pos`.

```python
indeed_feb2020_pos = ...
indeed_feb2020_pos.head()
```

```python
grader.check("q9")
```

Indeed (no pun intended) there weren't many days in February 2020 where the percent change in job postings was positive. To get a better since of what happened we can visualize each region using a line plot.

Run the cell below to import the `matplotlib` package.

```python
import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")
%matplotlib inline
```

<!-- BEGIN QUESTION -->

**Question 10.** Run the cell below to draw the line plot.

```python
# Get the values for each region
plt.plot(indeed_feb2020_pos['ral_cary'], label="Raleigh-Cary")
plt.plot(indeed_feb2020_pos['gso_hp'], label="Greensboro-High Point")
plt.plot(indeed_feb2020_pos['char_con_gast'], label="Charlotte-Concord-Gastonia")

# Add title
plt.title("Indeed Job Postings in February 2020")

# Add axes labels
plt.xlabel("Date")
plt.ylabel("Percent Change")

# Rotate x-tick marks 
plt.xticks(rotation=45)

# Add a legend
plt.legend();
```

<!-- END QUESTION -->

## Level IV 

Was there a major change a year later? Let's investigate.

**Question 11.** Select all the rows from the `indeed_date_idx` `DataFrame` for the month of February 2021. Store these rows in a `DataFrame` named `indeed_feb2021`.

```python
indeed_feb2021 = ...
indeed_feb2021.head()
```

```python
grader.check("q11")
```

<!-- BEGIN QUESTION -->

**Question 12.** Make a line plot that shows the percent change for February 2020 and February 2021 for the Raleigh Cary region.

**Hints:** 

* Use the code from **Question 10.** as a guide. 

* The number of days in February 2020 is 29, but the number of days in February 2021 is 28. You will need an array of values for February 2020 and another one for February 2021.

* Look at the documentation for the [range](https://www.w3schools.com/python/ref_func_range.asp) function to guide you on how to make an array of values.

```python
# Create an array of x-values for each year
...
...

# Get the values for the Raleigh Cary region for February in each year
plt.plot(...)
plt.plot(...)

# Add title
plt.title("Indeed Job Postings in February 2020 and 2021")

# Add axes labels
plt.xlabel("Day iof the Month")
plt.ylabel("Percent Change")

# Add a legend
plt.legend();
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 13.** Write a few sentences describing what you see in your graph from **Question 12.**.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Level V

### Wether Data

The NOAA (the US National Oceanic and Atmospheric Administration) operates weather stations that measure surface temperatures at different sites around the United States.  The hourly readings are [publicly available](http://www.ncdc.noaa.gov/qclcd/QCLCD?prior=N).

Suppose we download all the hourly data from the Raliegh-Durham, North Carolina site for the month of December 2021. To analyze the data, we want to know when each reading was taken, but we find that the data don't include the timestamps of the readings (the time at which each one was taken).

However, we know the first reading was taken at the first instant of December 2021 (midnight on December 1st) and each subsequent reading was taken exactly 1 hour after the last.

**Question 14.** Import `numpy` as `np` and then use `np.arange` to create an array of the *time, in seconds, since the start of the month* at which each hourly reading was taken.  Name it `collection_times`.

**Hints:** 

 - There were 31 days in December, which is equivalent to ($31 \times 24$) hours or ($31 \times 24 \times 60 \times 60$) seconds.  So your array should have $31 \times 24$ elements in it.

 - The `len` function works on arrays, too!  If your `collection_times` isn't passing the tests, check its length and make sure it has $31 \times 24$ elements.

```python
# Import the numpy package
import ... as np

# Create the collection times array
collection_times = ...
collection_times
```

```python
grader.check("q14")
```

### Uber

Below in the next two cells, we load dataframes containing 200,000 weekday Uber rides in the Manila, Philippines, and Boston, Massachusetts metropolitan areas from the [Uber Movement](https://movement.uber.com) project. The `sourceid` and `dstid` columns contain codes corresponding to start and end locations of each ride. The `hod` column contains codes corresponding to the hour of the day the ride took place. The `ride time` column contains the length of the ride, in minutes.

```python
boston = pd.read_csv('data/boston.csv')
print("Boston Table")
boston.head()
```

```python
manila = pd.read_csv('data/manila.csv')
print("Manila Table")
manila.head()
```

<!-- BEGIN QUESTION -->

**Question 15.** Create a histogram of all ride times in Boston using the given bins.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 16.** Create a histogram of all ride times in Manila using the given bins.

```python
...
```

<!-- END QUESTION -->

We can find the median ride time for each city in two ways. 

1. Using the `.median()` `Series` method.

2. Using the `NumPy` function `np.median()`.

The `pandas` `Series` object can be seen as an enhanced numpy 1-dimensional array and the `pandas` `DataFrame` can be seen as an enhanced numpy 2-dimensional array. The main difference is that `pandas` `Series` and `pandas` `DataFrames` has explicit index, while numpy arrays has implicit indexation. 

Run the next two cells.

```python
manila['ride time'].median()
```

```python
np.median(manila['ride time'])
```

**Question 17.** Find the average ride time for Manila and Boston.

```python
manila_avg = ...
boston_avg = ...
```

```python
grader.check("q17")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> 2022.0
2022.0
```


Tests for q11 (points: 0)
```
>>> isinstance(indeed_feb2021, pd.core.frame.DataFrame)
True
```
```
>>> indeed_feb2021.shape
(28, 3)
```


Tests for q14 (points: 0)
```
>>> isinstance(collection_times, np.ndarray)
True
```
```
>>> len(collection_times)
744
```
```
>>> all(collection_times == np.arange(0, 31*24*60*60, 60*60))
True
```


Tests for q17 (points: 0)
```
>>> import math
>>> math.isclose(boston_avg,  16.502477909999996)
True
```
```
>>> import math 
>>> math.isclose(manila_avg,  33.74324175500001)
True
```


Tests for q2 (points: None)
```
>>> seconds_in_a_decade == 315532800
True
```


Tests for q3 (points: None)
```
>>> min_height_difference == 0.03
True
```


Tests for q4 (points: None)
```
>>> near_twenty == 19.99909997918947
True
```


Tests for q5 (points: None)
```
>>> sine_of_pi_over_four == math.sin(math.pi/4)
True
```


Tests for q6 (points: None)
```
>>> isinstance(indeed, pd.core.frame.DataFrame)
True
```
```
>>> indeed.shape
(700, 4)
```
```
>>> indeed.columns
Index(['observation_date', 'ral_cary', 'char_con_gast', 'gso_hp'], dtype='object')
```


Tests for q7 (points: None)
```
>>> isinstance(indeed_date_idx, pd.core.frame.DataFrame)
True
```
```
>>> indeed_date_idx.shape
(700, 3)
```
```
>>> indeed_date_idx.columns
Index(['ral_cary', 'char_con_gast', 'gso_hp'], dtype='object')
```


Tests for q8 (points: None)
```
>>> isinstance(indeed_feb2020, pd.core.frame.DataFrame)
True
```
```
>>> indeed_feb2020.shape
(29, 3)
```


Tests for q9 (points: None)
```
>>> isinstance(indeed_feb2020_pos, pd.core.frame.DataFrame)
True
```
```
>>> indeed_feb2020_pos.shape
(7, 3)
```
