```python
# Initialize Otter
import otter
grader = otter.Notebook("week-2.ipynb")
```

# Week 2 Lecture Notebook

### Comments

Comments in Python are the lines in the code that are ignored by the compiler during the execution of the program. Comments enhance the readability of the code and help other programmers understand the code.

```python
# This is a comment.
# In this cell we adding evaluating an arithmetic expression.
3+5
```

```python
"""
This is another style of commenting using a string literal. 

In this cell we printing a message.
"""
print("Hello World")
```

### Names

Names are identifiers (names given to variables, functions, etc.) and keywords (reserved words). Variables are storage continers for data. Keywords can't be used as a variable names, function names, or any other identifiers. We will discuss keywords in more detail later.

```python
# variable names
# strings
course_name = "R/Python for Data Science"
course_code = "DSC495-001-009"

# integer
credit_hrs = 1

# float
gpa = 3.85
```

### Built-in Functions

The Python interpreter has a number of functions and types built into it that are always available. They are listed [here](https://docs.python.org/3/library/functions.html) in alphabetical order.

#### `print`

```python
print(course_name)
print(course_code)
print(credit_hrs)
print(gpa)
```

#### `type`

```python
# course_name and course_code are type string
type(course_name), type(course_code)
```

```python
# credit_hrs is type integer
type(credit_hrs)
```

```python
# gpa is type float (decimal)
type(gpa)
```

#### `round`

```python
round(gpa)
```

```python
help(round)
```

**Example 1.** Round `gpa` to one decimal place.

```python
...
```

```python
grader.check("e1")
```

#### `max`

```python
help(max)
```

**Example 2.** Return the maximum value between `credit_hrs` and `gpa`.

```python
...
```

```python
grader.check("e2")
```

### Iterable

An iterable is any Python object capable of returning its members one at a time, permitting it to be iterated over in a for-loop (we will discuss for-loops in more detail later). For instance, built-in types such as strings, lists, and dictionaries are all iterables. 

**Example 3.** What will be the output of `max(course_name)` and `min(course_code)`?

```python
max(course_name)
```

```python
min(course_code)
```

#### `list`

* Lists are used to store multiple items in a single variable.

* Lists are one of 4 built-in data types in Python used to store collections of data, the other 3 are tuple, set, and dictionary (we will discuss the dictionary in detail later).

* Lists are created using square brackets.

* Lists can contain multiple data types.

* Lists are mutable data types as the elements of the list can be modified. 

```python
my_list = ["Data", "Science", "Academy", "NC State", 495, 1, 9]
my_list
```

**Example 4.** How many elements are in `my_list`?

```python
...
```

```python
grader.check("e4")
```

We can access elements in a list by using the index values. In Python lists are one of the built-in data types used to store multiple items in a single variable. List items are ordered by index and the first index is `[0]`, the second item has index `[1]`, etc. 

```python
my_list[0]
```

**Example 5.** Change list element "NC State" to "NSCU".

```python
...
my_list
```

```python
grader.check("e5")
```

**Example 6.** Add NC State to `my_list`.

**Note:** The `.append()` method will add anitem to end of a list.

```python
...
```

```python
grader.check("e6")
```

```python
my_list
```

**Example 7.** Remove 495 from `my_list`. 

**Note:** There are multiple ways to remove items from a list

* Remove an item by value: `.remove()`

* Remove items by index or slice: `del`

```python
# The remove function works on the element name
...
```

```python
grader.check("e7")
```

```python
my_list
```

**Example 8.** Remove 1 from `my_list`.

```python
# The del function works on the element index
...
```

```python
grader.check("e8")
```

```python
my_list
```

### Import

The import command in Python is used to get access to other modules. A module is a set of functions and variables. To include or use these functions of modules in our code, we can simply import the module by using the `import` command.

#### `math`

The [`math`](https://docs.python.org/3/library/math.html) module provides access to the mathematical functions.

```python
import math
```

```python
math.factorial(5)
```

#### `NumPy`

[`NumPy`](https://numpy.org/doc/stable/user/index.html#user) is the fundamental package for scientific computing in Python.

```python
import numpy as np
```

```python
my_array = np.arange(11)
my_array
```

```python
max(my_array)
```

We can perform math operations to every element in a `NumPy` array.

```python
my_array**2
```

#### `pandas`

[`pandas`](https://pandas.pydata.org/docs/user_guide/10min.html) is a Python package providing fast, flexible, and expressive data structures designed to make working with “relational” or “labeled” data both easy and intuitive.

```python
import pandas as pd
```

```python
elections = pd.read_csv("data/elections.csv")

# The head function lists the first 5 rows of a pandas
# DataFrame by default
elections.head()
```

#### `shape`

[`shape`](https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.shape.html) returns a **tuple** representing the dimensionality of the DataFrame.

**Note:** A [tuple](https://www.w3schools.com/python/python_tuples.asp) is a collection of ordered items that is unchangeable. Tuples are written with round brackets.

```python
# The shape function returns the number of 
# rows and columns
elections.shape
```

#### `columns`

`columns` returns the column labels of the given Dataframe.

```python
# The columns function returns the
# names of the columns
elections.columns
```

We can look at the values of a spcific column using `[ ]`. This returns a `Series`. A `pandas` [`Series`](https://www.geeksforgeeks.org/python-pandas-series/) is like a column in a table. It is a one-dimensional array holding data of any type.

```python
# Returns a Series object
elections["%"]
```

To return a `DataFrame` use double brackets `[[ ]]`. The inner bracket is a list of column names from the `DataFrame`.

```python
# Returns a DataFrame object
elections[...]
```

There are two observations with `%` below 10. 

Who were these two candidates?

We can use `.loc` to find rows by index name. We will learn more about the index and using `.loc` next week. 

```python
# [2, 12] is a list of the rows we want from the DataFrame
elections.loc[..., :]
```

Click [here](https://www.sharpsightlabs.com/blog/pandas-index/) to a clear explanation of the `pandas` index.

Are there other Independent candidates? 

How well did they do?

We can use a boolean mask/array. Boolean masking is a way to quantify a sub-collection in a collection. Masking in Python and data science is when you want manipulated data in a collection based on some criteria. The criteria you use is typically of a true or false nature, hence the boolean part.

**Example 9.** Find all the rows where the `Party` is Independent?

```python
elections['...'] == '...'
```

```python
elections[elections['...'] == '...']
```

Let's visualize the distribution of the `%` column.

#### `matplotlib` 

Matplotlib is a comprehensive library for creating static, animated, and interactive visualizations in Python. Click [here](https://matplotlib.org/3.5.1/plot_types/index.html) to view plot types.

```python
import matplotlib.pyplot as plt
plt.style.use("fivethirtyeight")
%matplotlib inline
```

```python
# Hide gridlines
plt.grid(False)

# Add title
plt.title("Distributions of Vote Percentages")

# Add x label
plt.xlabel("Percent of Votes")

# Add y label
plt.ylabel("Frequency")

# Plot histogram 
plt.hist(elections["%"], edgecolor='black', linewidth=1);
```

Let's change the bin widths.

```python
# Set bin widths using a list
bin_widths = [0,20,40,60]

# Hide gridlines
plt.grid(False)

# Add title
plt.title("Distributions of Vote Percentages")

# Add x label
plt.xlabel("Percent of Votes")

# Add y label
plt.ylabel("Frequency")

# Plot histogram 
plt.hist(elections["%"], bins=bin_widths, edgecolor='black', linewidth=1);
```

Let's investigate another way a `DataFrame` can have an index.


```python
mottos = pd.read_csv('data/mottos.csv')
mottos.head()
```

By default the `mottos` `DataFrame` is indexed by integers. The `.set_index` method can be used to set a list, `Series`, or `DataFrame` as index of a `DataFrame`. Let's set the index of the `mottos` `DataFrame` to be the name of the state.

```python
mottos_state_idx = mottos.set_index('...')
```

```python
mottos_state_idx.head()
```

Now we can use the `.loc` method to locate observations by the name of the state.

```python
mottos_state_idx.loc['...']
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Tests for e1 (points: None)
```
>>> round(gpa, 1)
3.9
```


Tests for e2 (points: None)
```
>>> max(credit_hrs, gpa)
3.85
```


Tests for e4 (points: None)
```
>>> len(my_list)
7
```


Tests for e5 (points: None)
```
>>> my_list[3]
'NCSU'
```


Tests for e6 (points: None)
```
>>> len(my_list)
8
```
```
>>> my_list[7]
'NC State'
```


Tests for e7 (points: None)
```
>>> len(my_list)
7
```
```
>>> my_list[4]
1
```


Tests for e8 (points: None)
```
>>> len(my_list)
6
```
```
>>> my_list[4]
9
```
