```python
# Initialize Otter
import otter
grader = otter.Notebook("dsc495_0019_a1.ipynb")
```

# Assignment 1

## Due: Tuesday, January 18, 2022 @ 11:59pm

In this assignment we will get familiar with notebooks and some basic Python, including the print functions and variable assignment.

I would like you to attempt each level of the assignment. They get progressively harder. To get **full credit** on this assignment, you must complete **Level I** and **Level II**. **Level III** gives you a chance to work with some more advanced printing freatures in Python.

## Jupyter notebooks
This webpage is called a Jupyter notebook. A notebook is a place to write programs and view their results, and also to write text. For this class, I encourage the use of Notebooks to prototype work. In the future (either later in this class or beyond) you will likely want to write programs designed to run on all by themselves or interactively for end-users. But we generally will skip over that for now.

As for working in Notebooks, there are two types of cells (each rectangle containing text or code is called a **cell**): 

 - **code cells** that hold executable code.
 
 - **markdown cells** that hold a special kind of text that follows the markdown syntax. To get familiar with the markdown syntax, take a look at this [Markdown Gudie](https://www.markdownguide.org/basic-syntax/). Markdown cells (like this one) can be edited by double-clicking on them. After you edit a text cell, click the "run cell" button at the top that looks like ▶| or hold down `shift` + `return` to confirm any changes. 

   **Note:** Try not to delete the instructions of the assignment.

In the markdown cell below enter your name, section, and the date.

**Name:** 

**Section:** 

**Date:**

<!-- BEGIN QUESTION -->

## Level I

**Question 1.** Edit the code blocks below to answer each question with you own answer.

```python
# One thing to get used to: many programming languages start counting at element 0.
# So question 6 is actually the 7th question in the list above.
# It will take some time to get used to, expect to make off-by-one errors until this
# approach becomes second nature.
print("Question 0: What is your name?")
print("...")

print("\nQuestion 1: What motivated you to take a data science class?")
print("...")

print("\nQuestion 2: What do you hope to do once you graduate?")
print("...")

print("\nQuestion 3: What do you hope to learn in the class?")
print("...")

print("\nQuestion 4: Do you have any concerns about this class that you'd like to share?")
print("...")

print("\nQuestion 5: What is your prior programming experience?")
print("...")
```

<!-- END QUESTION -->

Notice how this print statement uses a special string character: `\n`. The leading slash `\` is special - it's called an escape charater. The following character tells python what to do with that escape charater. In our case, the `n` indicates there should be a new line before the rest of the text is printed. There are a number of special charaters you can find in Python. If you ever need to print a backslash, you can use it's own special character sequence: `\\`.

The backslash can also be used to continue a long line of code. Most standards recommend no more than 80 characters, including spaces, on a single line of code (although I use 88 since most modern monitors are big enough). I only suggest using a continuation backslash if you cannot find a place in your code to split the line neatly. We'll see example of that later in class.

## Level II

### Assignment 

The `=` character is not an equal sign in python. It's used for **assignment**. The name on the left is being assigned the object on the right. In our case, we want to assign a string (a group of characters contained inside quotation marks") to some variable name.

### Variable Names 

Variable names should be informative so that you know what they mean when you come back to your code. For example, you might forget what `answer_0` means, and without good commenting someone else reading you code will likely be confused too, but `my_name` gives a pretty good indication of what that variable holds.

There are also some common conventions on how to name variables. For now, we will primarily use lower case words separated by underscores.

### Checking your code
Now that you know how to name things, you can start using the built-in tests to check whether your work is correct. Sometimes, there are multiple tests for a single question, and passing all of them is required to receive credit for the question. A test cell will always look like:

```
grader.check("q1")
```

Go ahead and attempt **Question 2.** Running the test cell directly after it will test whether you have assigned `my_name` correctly. Sometimes if you make a common error the tests will give hints about what went wrong.

Let's practice. 

**Question 2.** Assign your name to the variable `my_name` by replacing the ellipsis (`...`) with your name in quotation marks, then print your answer to the question. Here's an example:

```
my_name = "Mahmoud Harding"
print("Question 0: What is your name?")
print(my_name)
```

You can also put the variable and other text string together in a print function separated by commas, like this:
```
print("Question 0: What is your name?\n", my_name)
```

```python
my_name = ...
print("Question 0. What is your name?")
print(my_name)
```

```python
grader.check("q2")
```

### Errors

Python is a language, and like natural human languages, it has rules.  It differs from natural language in two important ways:

1. The rules are *simple*.  You can learn most of them in a few weeks and gain reasonable proficiency with the language in a semester.

2. The rules are *rigid*.  If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes.  A computer running Python code is not smart enough to do that.

Whenever you write code, you'll make mistakes.  When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong. Errors are okay; even experienced programmers make many errors.  When you make an error, you just have to find the source of the problem, fix it, and move on.

There is an error in the next cell.  Run it and see what happens.

```python
print("This line is missing something."
```

**Note:** In the toolbar, there is the option to click `Run > Run All Cells`, which will run all the code cells in this notebook in order. However, the notebook stops running code cells if it hits an error, like the one in the cell above.

You should see something like this (minus our annotations):

<img src="images/error.jpg" />

The last line of the error output attempts to tell you what went wrong.  The *syntax* of a language is its structure, and this `SyntaxError` tells you that you have created an illegal structure.  "`EOF`" means "end of file," so the message is saying Python expected you to write something more (in this case, a right parenthesis) before finishing the cell.

There's a lot of terminology in programming languages, but you don't need to know it all in order to program effectively. If you see a cryptic message like this, you can often get by without deciphering it. Of course, if you're frustrated, ask a classmate, a TA, or your instructor for help.

Try to fix the code below so that you can run the cell and see the intended message instead of an error.

```python
print("This line is missing something."
```

## Level III

Python offers a variety of ways to print, so let's explore two of them. They are the `.format()` and `f-string` methods. Both of these methods give you a greater degree of control over print formatting, especially when dealing with a lot of variables in a single print statement. 

What these methods have in common is the use of curly brackets `{}` used as placeholders where a variable will be printed inside the string.

With `.format()`, the curly brackets in the string are empty, and the format function holds the variable(s) that are to be printed at each curly bracket. If the number of variables in the function does *not* match the number of curly brackets, an error will be given.

```
my_name = "Mahmoud"
print("Question 0: What is your name?\n{}".format(my_name))
```

The `f-string` approach requires us to put the letter f before the opening quotation mark. The actual variable name is then put inside the curly brackets.

```
my_name = "Mahmoud"
print(f"Question 0: What is your name?\n{my_name}")
```

**Question 3.** Copy your code block from **Question 2.** where you assign you answers to variables and practice printing your answers with `.format()`.

```python
my_name = "..."
print("Question 0. What is your name?\n{}".format(...))
```

```python
grader.check("q3")
```

**Question 4.** Copy your code block from **Question 2.** where you assign you answers to variables and practice printing your answers with `f-strings`.

```python
my_name = "..."
print(f"Question 0: What is your name?\n{...}")
```

```python
grader.check("q4")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q2 (points: None)
```
>>> type(my_name) == str
True
```


Tests for q3 (points: None)
```
>>> type(my_name) == str
True
```


Tests for q4 (points: None)
```
>>> isinstance(my_name, str)
True
```
